---
title: 'Hot-deck Imputation'
description: 'Learner will be able to impute values based on how similar those values are to other values using R. '
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Impute by Matching

```yaml
type: VideoExercise
key: a297bfdc4c
xp: 50
```

`@projector_key`
8a75340ffe779b8d3cfdd8597519ffe0
